# Module 5. Advanced functions.

**_Author: Favio Vázquez_**

**Expected time = 2.5 hours**

**Total points = 130 points**


## Assignment Overview

In this assignment you will work with Python functions, from the basics of their structure to more nuanced activities. You will start by reviewing what you learned about functions in earlier modules, then you will practice how to cast arguments inside of functions, how to use the local and the global environment, how to use nested functions, and then how to deal with exceptions. In the final part of the assignment you will be testing your knowledge of lambda functions and the basics of functional programming.

This assignment is designed to build your familiarity and comfort coding in Python while also helping you review key topics from each module. As you progress through the assignment, answers will get increasingly complex. It is important that you adopt a data scientist's mindset when completing this assignment. **Remember to run your code from each cell before submitting your assignment.** Running your code beforehand will notify you of errors and give you a chance to fix your errors before submitting. You should view your Vocareum submission as if you are delivering a final project to your manager or client. 

***Vocareum Tips***
- Do not add arguments or options to functions unless you are specifically asked to. This will cause an error in Vocareum.
- Do not use a library unless you are explibitly asked to in the question. 
- You can download the Grading Report after submitting the assignment. This will include feedback and hints on incorrect questions.


### Learning Objectives

- Construct complex functions including functions with multiple parameters, nested functions, and functions with default vs. flexible arguments. 
- Distinguish between global, local, and built-in scope. 
- Interpret errors and understand exceptions in order to troubleshoot.  
- Use lambda functions to analyze DataFrames.

## Index:

#### Module 5: Advanced functions

- [Question 1](#Question-1)
- [Question 2](#Question-2)
- [Question 3](#Question-3)
- [Question 4](#Question-4)
- [Question 5](#Question-5)
- [Question 6](#Question-6)
- [Question 7](#Question-7)
- [Question 8](#Question-8)
- [Question 9](#Question-9)
- [Question 10](#Question-10)
- [Question 11](#Question-11)
- [Question 12](#Question-12)


## Module 5: Advanced functions

In the first part of this assignment, we will reviewing the basics of functions and the way we create them in Python.

### Basic functions review

We will begin this assignment with a review of functions to refresh your memory on the basics. As you may remember, functions are followed by parentheses () and the argument is within the parentheses like this: `function(argument)` . The arguments are the input that the function performs some calculations on, and then produces an output. You may use either built-in functions or create your own user-defined functions. Functions may also have multiple parameters. 

[Back to top](#Index:) 

### Question 1
*5 points*

Create a function called `square_cond` that takes an argument and returns the square of it, if and only if the number is even. If the number is odd return the sentence "Odd number". 

```python
### GRADED

### YOUR SOLUTION HERE
def square_cond():
    return

### BEGIN SOLUTION
def square_cond(x):
    if x % 2 == 0:
        return x*x
    else:
        return "Odd number"
### END SOLUTION
```

<!-- grade_id: Question 01, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def square_cond_(x):
    if x % 2 == 0:
        return x*x
    else:
        return "Odd number"
#
#
#
assert square_cond(10) == 100
assert square_cond(11) == "Odd number"
assert square_cond_(1) == square_cond(1)
assert square_cond_(2) == square_cond(2)
assert square_cond_(14) == square_cond(14), "Did you compute the square of the number correctly?"
assert square_cond_(7) == square_cond(7), "Did you make sure your function only takes even numbers?"
print("That's correct!")
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Question 2
*15 points*

Create a Python function called `multi_operation` that takes two floats `a` and `b` and returns a tuple. Keep the following in mind:
- Ensure that both `a` and `b` are floats. If not they are not floats, parse them inside the function. 
- The first element of the tuple must be the multiplication of `a` and `b`. 
- The second element must be the integer division of `a` and `b`. 
- Return both elements.

**Hint: The integer division is perfomed by the `//` operator.**

```python
### GRADED

### YOUR SOLUTION HERE
def multi_operation():
    return

### BEGIN SOLUTION
def multi_operation(a,b):
    if not isinstance(a,float):
        a = float(a)
    if not isinstance(b,float):
        b = float(b)
    multi = a*b
    div = a//b
    return multi,div
### END SOLUTION
```

<!-- grade_id: Question 02, points: 15, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def multi_operation_(a,b):
    if not isinstance(a,float):
        a = float(a)
    if not isinstance(b,float):
        b = float(b)
    multi = a*b
    div = a//b
    return multi,div
#
#
#
assert multi_operation(1,2) == (2.0, 0.0)
assert multi_operation(10.2,2.2) == (22.44, 4.0)
assert multi_operation(20,10) == multi_operation_(20,10)
assert multi_operation(40,2.5) == multi_operation_(40,2.5), "Did you make sure that your function only takes floats?"
print("That's correct!")
### END HIDDEN TESTS
```

### Local vs Global scope

In this section you will test your knowledge about scope in Python. As you may remember: 
- Global scope is defined in the main body of the script.
- Local scope is defined within a function.
- Built-in scope comes standard with Python.

[Back to top](#Index:) 

### Question 3
*10 points*

Remember that variables that are defined inside a function body have a local scope, and those defined outside of the function body have a global scope.With that in mind, complete the following: 
- Define a function called `global_sum` that takes two integers and computes their sum. 
- Create a global variable, inside the function, named `number` to store the result without returning it.

```python
### GRADED

### YOUR SOLUTION HERE
def global_sum():
    return 

### BEGIN SOLUTION
def global_sum(a,b):
    global number
    number = a+b
### END SOLUTION
```

<!-- grade_id: Question 03, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
#
#
#
global_sum(10,4)
assert number == 14, "Did you define number as a global variable?"
global_sum(2,23)
assert number == 25
global_sum(-5,4)
assert number == -1
global_sum(223,-1233)
assert number == -1010
print("That's correct!")
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Question 4
*5 points*

Create a function called `i_print` that takes no argument and returns the global variable `x` as an integer. A few items to note:
- `x` is already defined for you. 
- The printed value must be an integer.

```python
### GRADED

### YOUR SOLUTION HERE
# Definition of x
x = 10

def i_print():
    return 

### BEGIN SOLUTION
def i_print():
    return int(x)
### END SOLUTION
```

<!-- grade_id: Question 04, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
x_ = 10

def i_print_():
    return int(x_)
#
#
#
assert i_print() == i_print_(), "Did you cast x to an integer?"
print("That's correct!")
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Question 5

*5 points*

We define a function as:

```python

a = 10

def h():     
    global a 
    a = 3
    return a
```

What will the function `h()` return? Store your solution in the variable `ans_1`.

```python
### GRADED

### YOUR SOLUTION HERE
ans_1 = None

### BEGIN SOLUTION
ans_1 = 3
### END SOLUTION
```

<!-- grade_id: Question 05, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert ans_1 == 3, "Are you sure?"
print("That's correct!")
### END HIDDEN TESTS
```

### Nested functions and functions with multiple arguments

In this section you will test your knowledge on how to create and use nested functions and how to define functions with default parameters.

We use nested functions to run a process multiple times. You may remember:
- A nested function is a function within a function.
- Return is used to return the variable so it can be used by the function.
- The innermost function is processed first.
- The outermost function is processed last.

We can use built-in functions with default arguments and specify default arguments for our user-defined functions. 
- Default arguments make it optional for the user to enter an argument.
- We can use an asterisk before our parameter name to indicate that any number of arguments can be passed through the function and use a for loop so the function applies the code to each argument.

[Back to top](#Index:) 

### Question 6
*15 points*

Define a nested function called `nested_sum`, where in the first part of the function you accept an argument called `x` and in the second part (the function `f`  that is inside) you take another argument called `y`. 
- In the function `f`, calculate the sum of `x` and `y`. 
- To test your function, create a variable called `res_1` where you pass the `x` argument to `nested_sum`, and then create a variable called `res_2` where you pass the `y` argument of the res_1 variable to get the final solution. 
- Have `x` equal 2 for res_1 and `y` equal 10 for res_2.

```python
### GRADED

### YOUR SOLUTION HERE
def nested_sum():
    return

res_1 = None
res_2 = None

### BEGIN SOLUTION
def nested_sum(x):
    def f(y):
        return x+y
    return f

res_1 = nested_sum(2)
res_2 = res_1(10)
### END SOLUTION
```

<!-- grade_id: Question 06, points: 15, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def nested_sum_(x):
    def f_(y):
        return x+y
    return f_

res_1_ = nested_sum_(2)
res_2_ = res_1_(10)

res_3 = nested_sum(10)
res_4 = res_3(10)

res_5 = nested_sum_(10)
res_6 = res_5(10)
#
#
#
assert res_2 == res_2_, "Have you defined your nested functions correctly?"
assert res_4 == res_6
print("That's correct!")
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Question 7
*10 points*

Define a function called `i_divide` that takes two arguments: `num_1` and `num_2`, where `num_2` is 2 by default. The return statement should divide `num_1` by `num_2`. 

```python
### GRADED

### YOUR SOLUTION HERE
def i_divide():
    return

### BEGIN SOLUTION
def i_divide(num_1, num_2=2):
    return num_1 / num_2
### END SOLUTION
```

<!-- grade_id: Question 07, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def i_divide_(num_1, num_2=2):
    return num_1 / num_2
#
#
#
assert i_divide(10,2) == i_divide_(10,2) , "Did you set the second default argument correctly?"
assert i_divide(1,22) == i_divide_(1,22), "Are you performing the correct math operation?"
assert i_divide(num_2=10, num_1=10) == i_divide_(num_2=10, num_1=10)
print("That's correct!")
### END HIDDEN TESTS
```

### Handling exceptions

In this section you will test your knowledge on how to handle different types of exceptions and raise errors in Python.
- You can use `try` and `except` to print a customized error message. 
- Remember to specify the error type after `except` to ensure that error messages are specific to potential errors. 
- You can use an if statement to create more refined error messages. 

[Back to top](#Index:) 

### Question 8
*20 points*

Define a function called `exce_sum` where you return the sum of two arguments, but if the both arguments are 0, the function should raise an exception saying "Invalid numbers". The exception needs to have the ValueError() class and the return type of the exception must be a string.

**Hint: Make sure to validate if the two arguments are 0. Use `try` and `except`.**

```python
### GRADED

### YOUR SOLUTION HERE
def exce_sum():
    return

### BEGIN SOLUTION
def exce_sum(a,b):
    try:
        if a == 0 and b==0:
            raise ValueError("Invalid numbers")
        return a+b
    except ValueError as err:
        return str(err)
### END SOLUTION
```

<!-- grade_id: Question 08, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def exce_sum_(a,b):
    try:
        if a == 0 and b==0:
            raise ValueError("Invalid numbers")
        return a+b
    except ValueError as err:
        return str(err)
#
#
#
assert exce_sum(1,2) == exce_sum_(1,2), "Did you the raise the exception correctly?"
assert exce_sum(1,1) == exce_sum_(1,1)
assert exce_sum(0,1) == exce_sum_(0,1)
assert exce_sum(0,0) == "Invalid numbers"
assert exce_sum(0,0) == exce_sum_(0,0)
assert exce_sum(1,6) == exce_sum_(1,6)
print("That's correct!")
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Question 9
*15 points*

Define a function called `exce_sub` that takes two arguments and returns the difference between two arguments. If both arguments are 1, the function should raise an exception saying "Invalid numbers". The exception must be inherited from the ValueError() class and the return type of the exception must be a string.

**Hint: Make sure to validate if the two arguments are 0. Use `try` and `except`.**

```python
### GRADED

### YOUR SOLUTION HERE
def exce_sub():
    return

### BEGIN SOLUTION
def exce_sub(a,b):
    try:
        if a == 1 and b==1:
            raise ValueError("Invalid numbers")
        return a-b
    except ValueError as err:
        return str(err)
### END SOLUTION
```

<!-- grade_id: Question 09, points: 15, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def exce_sub_(a,b):
    try:
        if a == 1 and b == 1:
            raise ValueError("Invalid numbers")
        return a-b
    except ValueError as err:
        return str(err)
#
#
#
assert exce_sub(2,1) == exce_sub_(2,1), "Did you the raise the exception correctly?"
assert exce_sub(10,3) == exce_sub_(10,3)
assert exce_sub(0,1) == exce_sub_(0,1)
assert exce_sub(1,1) == "Invalid numbers"
assert exce_sub(1,1) == exce_sub_(1,1)
assert exce_sub(1,6) == exce_sub_(1,6)
print("That's correct!") 
### END HIDDEN TESTS
```

### Lambda functions

In this final section you will test your knowledge of lambda functions. Remember, you can use the `lambda` keyword to write functions quickly, but only for functions that evaluate to an expression. So this is a good approach for functions you will only use a few times. 

As you may remember, lambda functions look like this. 

`name = lambda argument: what the function will do`

[Back to top](#Index:) 

### Question 10
*5 points*
  
Create a variable called `f` that stores a lambda function to square a number. Test your function with the number 6 and store it in a variable called `ans_1`, then with the number 112 and store it in a variable called `ans_2`.

```python
### GRADED

### YOUR SOLUTION HERE
f = None

ans_1 = None
ans_2 = None

### BEGIN SOLUTION
f = lambda x: x*x

ans_1 = f(6)
ans_2 = f(112)
### END SOLUTION
```

<!-- grade_id: Question 10, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
f_ =  lambda x: x*x

ans_1_ = f_(6)
ans_2_ = f_(112)

#
#
#
assert ans_1 == ans_1_, "Did you define your lambda function correctly?"
assert ans_2 == ans_2
assert f(10) == 100
assert f(4) == 16
print("That's correct!")
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Question 11
*10 points*
    
Create a list called `res` that stores a lambda function using the `map` function to capitalize all the strings inside of the list `try_list`.

```python
### GRADED

### YOUR SOLUTION HERE
try_list = ["programming", "is", "awesome", "and", "more", "with", "emeritus"]
res = None

### BEGIN SOLUTION
res = list(map(lambda x: x.capitalize(), try_list))
### END SOLUTION
```

<!-- grade_id: Question 11, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
try_list_ = ["programming", "is", "awesome", "and", "more", "with", "emeritus"]
res_ = list(map(lambda x: x.capitalize(), try_list_))
#
#
#
assert try_list == try_list_, "Make sure you use the correct string method here"
assert res == res_
print("That's correct!")
### END HIDDEN TESTS
```

[Back to top](#Index:)

### Question 12
*15 points*

Given the dataframe `df`, define a function called `multi` that takes two arguments `x` and `y` that returns `x` multiplied by `y`. Use a lambda function inside the `df.apply` function to multiply the columns `num_legs` and `num_specimen` utilizing the multi function you defined, and then store the final result in a new column called `result` in the same dataframe.

**Hint: Use `df.apply`**

```python
### GRADED

### YOUR SOLUTION HERE

import pandas as pd

df = pd.DataFrame({'type': ['falcon', 'dog', 'spider', 'fish'],
                   'num_legs': [2, 4, 8, 0],
                   'num_wings': [2, 0, 0, 0],
                   'num_specimen': [10, 2, 1, 8]})

df["result"] = None

### BEGIN SOLUTION

def multi(x,y):
    return x*y

df["result"] = df.apply(lambda x: multi(x["num_legs"],x["num_specimen"]),axis=1)
### END SOLUTION
```

<!-- grade_id: Question 12, points: 15, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from pandas.util.testing import assert_frame_equal
import pandas as pd

df_ = pd.DataFrame({'type': ['falcon', 'dog', 'spider', 'fish'],
                   'num_legs': [2, 4, 8, 0],
                   'num_wings': [2, 0, 0, 0],
                   'num_specimen': [10, 2, 1, 8]})

def multi(x,y):
    return x*y

df_["result"] = df_.apply(lambda x: multi(x["num_legs"],x["num_specimen"]),axis=1)
#
#
#
assert df.equals(df_), "Did you define the function `multi` correctly?"
print("That's correct!")
### END HIDDEN TESTS
```