# %title%

**_Author: Favio Vázquez_**

**_Reviewer: Jessica Cervi_**

**Expected time = %expected_time%**

**Total points = 125 points**


## Assignment Overview

In this assignment you will work with Python functions, from the basics of their structure to more nuanced activities. You will start by reviewing what you learned about functions in earlier modules, then you will practice how to cast arguments inside of functions, how to use the local and the global environment, how to use nested functions, and then how to deal with exceptions. In the final part of the assignment you will be testing your knowledge of lambda functions and the basics of functional programming.

This assignment is designed to build your familiarity and comfort coding in Python while also helping you review key topics from each module. As you progress through the assignment, answers will get increasingly complex. It is important that you adopt a data scientist's mindset when completing this assignment. **Remember to run your code from each cell before submitting your assignment.** Running your code beforehand will notify you of errors and give you a chance to fix your errors before submitting. You should view your Vocareum submission as if you are delivering a final project to your manager or client. 

***Vocareum Tips***
- Do not add arguments or options to functions unless you are specifically asked to. This will cause an error in Vocareum.
- Do not use a library unless you are explibitly asked to in the question. 
- You can download the Grading Report after submitting the assignment. This will include feedback and hints on incorrect questions.


### Learning Objectives

- Construct complex functions including functions with multiple parameters, nested functions, and functions with default vs. flexible arguments. 
- Distinguish between global, local, and built-in scope. 
- Interpret errors and understand exceptions in order to troubleshoot.  
- Use lambda functions to analyze DataFrames.

## Index:

#### %title%

- [Question 1](#Question-1)
- [Question 2](#Question-2)
- [Question 3](#Question-3)
- [Question 4](#Question-4)
- [Question 5](#Question-5)
- [Question 6](#Question-6)
- [Question 7](#Question-7)
- [Question 8](#Question-8)
- [Question 9](#Question-9)
- [Question 10](#Question-10)
- [Question 11](#Question-11)
- [Question 12](#Question-12)


## %title%

In the first part of this assignment, we will be reviewing the basics of functions and the way we create them in Python.

### Basic functions review

We will begin this assignment with a review of how to write functions to refresh your memory of the basics. Every function definition in Python starts with the keyword `def` followed by the name of the function.
The name of the function is followed by parentheses that are suppose to include *arguments*, if there are any.

The header of the function is followed by the *body of the function*. This is meant to be an **indented**
block of code containing the instructions that the function is supposed to perform.

Sometime a function ends with a `return` statement. A return statement is used to end the execution of the function call and 'returns' the result (value of the expression following the return keyword) to the caller. 

The pseudocode below displays the general structure of a function.
 
 
 ``` Python
def name_of_function(args):
    body of the function
    return var
```


Before starting the assignment, we will provide you with an example function to practice with.

### Example 1

Define a function called `div_by_five` that takes a number `x` as input. Your function should check whether `x` is divisible by five and, if so, return `x` cubed. Otherwise, your function should return the string `Number not divisible by five`.


### Solution

Observe the code cell below, where we have defined the function according to the instructions given above:

```python
def div_by_five(x):
    if (x % 5 == 0):
        return x**3
    else:
        return 'Number not divisible by five'
```

In the function header, we have defined the function name `div_by_five` and listed the arguments that our function takes. Next, we checked, using an `if` statement, whether x is divisible by five or not. If yes, the function returns the cube of `x`, otherwise we use `else` to return the string `Number not divisible by five`.


**IMPORTANT NOTE: In some questions, we will ask you to assign the result of a certain operation to a variable and then return that variable. Make sure to read the instructions for each question carefully.**

Before starting the assignment, feel free to change the argument passed to `div_by_five` in the code cell below to see the function in action.

```python
div_by_five(7)
```

Output:
```
'Number not divisible by five'
```

[Back to top](#Index:) 
### Question 1
*5 points*

Define a function called `square_cond` that takes takes the variable x as input, which should be a 
number.  Your function should check whether `x` is even  and, if so, returns `x` squared. Otherwise, your function should return the string "Odd number".


```python
### GRADED

### YOUR SOLUTION HERE
def square_cond():
    return

### BEGIN SOLUTION
def square_cond(x):
    if x % 2 == 0:
        return x*x
    else:
        return "Odd number"
### END SOLUTION
```

<!-- grade_id: Question 01, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (5)
def square_cond_(x):
    if x % 2 == 0:
        return x*x
    else:
        return "Odd number"
#
#
#
assert square_cond(10) == 100, "Did you compute the square of the number correctly?"
assert square_cond_(2) == square_cond(2), "Did you compute the square of the number correctly?"
assert square_cond_(14) == square_cond(14), "Did you compute the square of the number correctly?"
assert type(square_cond(7)) is str, "Did you make sure your function only takes even numbers?"
assert square_cond(11) == square_cond_(11), "Are you returning the correct string on even numbers?"
assert square_cond(1) == square_cond_(1), "Are you returning the correct string on even numbers?"
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

[Back to top](#Index:) 

### Question 2
*15 points*

Functions can also accept multiple arguments and return multiple variables at once by separating them with a comma.

Create a Python function called `multi_operation` that takes, as input, two floats, `a` and `b`.

Your function should compute the product of `a` and `b` and assign the result to `multi`. Next, your function should compute the integer division between `a` and `b` and assign the result to `div`.
The function should return both the `multi` and the `div` variables.

**Hint: The integer division is perfomed by the `//` operator.**

```python
### GRADED

### YOUR SOLUTION HERE
def multi_operation():
    return

### BEGIN SOLUTION
def multi_operation(a,b):
    multi = a * b
    div = a // b
    return multi, div
### END SOLUTION
```

<!-- grade_id: Question 02, points: 15, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (15)
def multi_operation_(a,b):
    multi = a * b
    div = a // b
    return multi, div
#
#
#
import math
assert len(multi_operation(1, 2)) == 2, "Did you return both variables?"
assert multi_operation(1, 2)[0] == multi_operation_(1, 2)[0], "Is your multiplication correct?"
assert multi_operation(1, 2)[1] == multi_operation_(1, 2)[1], "Is your integer division correct?"
assert math.isclose(multi_operation(10.2, 2.2)[0], multi_operation_(10.2, 2.2)[0]), "Is your multiplication correct?"
assert math.isclose(multi_operation(10.2, 2.2)[1], multi_operation_(10.2, 2.2)[1]), "Is your integer division correct?"
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

### Local vs Global scope

In this section you will test your knowledge about the different scopes in Python. Remember,

- Variables have global scope in the *main body* of the script
- Variables defined within a function have local scope by default
- We can define global variables within a function by using the keywird global

Observe the code below:

```Python

a = "Data science is awesome!"

def my_fun(x)
    global result
    result = x**3
    result_squared = result**2
    
```

In the code above, the variable `a` is defined in the main body. Therefore `a` can be used *everywhere* in my code.

In the function `my_fun`, we have defined the variable `result` to be global. This means that, after calling the function, we can access the value of `result` from the main body of my code without returning the variable with the function. On the other hand, `result_squared` is defined locally inside the function, therefore I can only use that variable within `my_fun`.

[Back to top](#Index:) 

### Question 3
*10 points*

Remember that variables that are defined inside a function body have a local scope, and those defined outside have a global scope. However, you can make variables inside a function global by declaring them using the `global` keyword.

Define a function called `global_sum` that takes two integers `a` and `b`. 
Inside the function, create a global variable `number` to store the result.
Assign to `number` the result of the sum between `a` and `b`.

**NOTE: Because `number` is defined globally, we can access its values without a return statement.**

```python
### GRADED

### YOUR SOLUTION HERE
def global_sum():
    return 

### BEGIN SOLUTION
def global_sum(a, b):
    global number
    number = a + b
### END SOLUTION
```

<!-- grade_id: Question 03, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (10)
def global_sum_(a, b):
    global number_
    number_ = a + b
#
#
#
assert 'global_sum' in globals(), "Did you define the 'global_sum' function?"
global_sum(10,4)
global_sum_(10,4)
assert 'number' in globals(), "Did you define 'number' as a global variable?"
assert number == number_, "Did you define 'number' as a global variable?"
del number, number_
global_sum(2,23)
global_sum_(2,23)
assert 'number' in globals(), "Did you define 'number' as a global variable?"
assert number == number_,"Did you define 'number' as a global variable?"
del number, number_
global_sum(-5,4)
global_sum_(-5,4)
assert 'number' in globals(), "Did you define 'number' as a global variable?"
assert number == number_,"Did you define 'number' as a global variable?"
del number, number_
global_sum(223,-1233)
global_sum_(223,-1233)
assert 'number' in globals(), "Did you define 'number' as a global variable?"
assert number == number_
del number, number_
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

[Back to top](#Index:) 

### Question 4
*5 points*

In the code cell below, we have defined the variable `x` for you.

Create a function called `i_print` that takes one argument `x`. The function should return the value `x` as an integer.

```python
### GRADED

### YOUR SOLUTION HERE
# Definition of x
x = 10.0

def i_print():
    return 

### BEGIN SOLUTION
def i_print(x):
    return int(x)
### END SOLUTION
```

<!-- grade_id: Question 04, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (5)
x_ = 10.0

def i_print_(x_):
    return int(x_)
#
#
#
assert i_print(x) == i_print_(x_), "Did you have the function return the correct value for x?"
assert type(i_print(x)) == type(i_print_(x_)), "Did you cast x to an integer?"
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

[Back to top](#Index:) 

### Question 5

*5 points*

We define a function as:

```python

a = 10

def h():     
    global a 
    a = 3
    return a
```

What will the function `h()` return? Store your solution in the variable `ans_1`.

```python
### GRADED

### YOUR SOLUTION HERE
ans_1 = None

### BEGIN SOLUTION
ans_1 = 3
### END SOLUTION
```

<!-- grade_id: Question 05, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (5)
assert ans_1 == 3, "Are you sure?"
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

### Nested functions and multiple arguments

In this section you will test your knowledge on how to create and use nested functions and how to define functions with default parameters.

We use nested functions to run a process multiple times. You may remember:
- A nested function is a function within a function.
- Return is used to return the variable so it can be used by a function.
- The innermost function is processed first.
- The outermost function is processed last.

We can use built-in functions with default arguments and specify default arguments for our user-defined functions. 
- Default arguments make it optional for the user to enter an argument.
- We can use an asterisk before our parameter name to indicate that any number of arguments can be passed through the function and use a for loop so the function applies the code to each argument.

[Back to top](#Index:) 

### Question 6
*15 points*

Define a function called `nested_sum`. This function is going to be the ourtermost one of your nested functions and should accept an argument  `x`.

Within `nested_sum`, define a function `f` that accepts an argument `y`. The function `f` should calculate and return the sum of `x` and `y`. 

Have `nested_sum` return `f`.

To test your function create a variable called `res_1` where you pass the  argument `x=2` to `nested_sum`, and then create a variable called `res_2` where you pass the argument `y = 10` of the `res_1` variable to get the final solution. 


```python
### GRADED

### YOUR SOLUTION HERE
def nested_sum():
    return

res_1 = None
res_2 = None

### BEGIN SOLUTION
def nested_sum(x):
    def f(y):
        return x+y
    return f

res_1 = nested_sum(x=2)
res_2 = res_1(y=10)
### END SOLUTION
```

<!-- grade_id: Question 06, points: 15, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (15)
def nested_sum_(x):
    def f_(y):
        return x+y
    return f_

res_1_ = nested_sum_(2)
res_2_ = res_1_(10)

res_3 = nested_sum(10)
res_4 = res_3(10)

res_5 = nested_sum_(10)
res_6 = res_5(10)
#
#
#
assert res_2 == res_2_, "Have you defined your nested functions correctly?"
assert res_4 == res_6, "Have you defined your nested functions correctly?"
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

[Back to top](#Index:) 

### Question 7
*10 points*

Define a function called `i_divide` that takes two arguments: `num_1` and `num_2`, where `num_2` equals 2 *by default*. The return statement should divide `num_1` by `num_2`.

```python
### GRADED

### YOUR SOLUTION HERE
def i_divide():
    return

### BEGIN SOLUTION
def i_divide(num_1, num_2=2):
    return num_1 / num_2
### END SOLUTION
```

<!-- grade_id: Question 07, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (10)
def i_divide_(num_1, num_2=2):
    return num_1 / num_2
#
#
#
import math
assert math.isclose(i_divide(10), i_divide_(10)) , "Did you set the second default argument correctly?"
assert math.isclose(i_divide(1,22), i_divide_(1,22)), "Are you performing the correct math operation?"
assert math.isclose(i_divide(num_2=10, num_1=10), i_divide_(num_2=10, num_1=10)), "Are you performing the correct math operation?"
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

### Handling exceptions

In this section you will test your knowledge on how to handle different types of exceptions and raise errors in Python.
- You can use `try` and `except` to print a customized error message. 
- Remember to specify the error type after `except` to ensure that error messages are specific to potential errors. 
- You can use an if statement to create more refined error messages. 

[Back to top](#Index:) 

### Question 8
*20 points*

Define a function called `exce_sum` that takes two arguments, `a` and `b`.

If both arguments are  0, the function should raise an `ValueError` exception saying "Invalid numbers". Otherwise, your function should return the sum of two arguments.

Please review the psuedocode for the function to help you get started:

```Python
def exce_sum(arguments):
    try:
        if both numbers are zero
            raise ValueError("Some error message")
        return  sum of arguments
    except ValueError as err
        return str(err)
```

```python
### GRADED

### YOUR SOLUTION HERE
def exce_sum():
    return

### BEGIN SOLUTION
def exce_sum(a,  b):
    if a == 0 and b == 0:
        raise ValueError("Invalid numbers")
    return a + b
### END SOLUTION
```

<!-- grade_id: Question 08, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (20)
def exce_sum_(a, b):
        if a == 0 and b == 0:
            raise ValueError("Invalid numbers")
        return a + b
#
#
assert exce_sum(1, 2) == exce_sum_(1, 2), "Did you compute the sum correctly?"
assert exce_sum(1, 1) == exce_sum_(1, 1), "Did you compute the sum correctly?"
assert exce_sum(1, 0) == exce_sum_(1, 0), "Did you compute the sum correctly?"
try:
    exce_sum(0, 0)
except Exception as e:
    assert isinstance(e, ValueError), "Make sure to raise a ValueError and not any other type of Exception."
    assert str(e) == "Invalid numbers", "Your error message '{}' is not correct.".format(e)
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

[Back to top](#Index:) 

### Question 9
*15 points*

Define a function called `exce_sub` that takes two arguments, `a` and `b`.

If both arguments are  1, the function should raise a `ValueError` exception saying "Invalid numbers". Otherwise, your function should return the difference of `a` and `b`.

Please review the psuedocode for the function to help you get started:

```Python
def exce_sub(arguments):
    if both numbers are one
        raise ValueError("Some error message")
    return difference of arguments
```

```python
### GRADED

### YOUR SOLUTION HERE
def exce_sub():
    return

### BEGIN SOLUTION
def exce_sub(a,b):
    if a == 1 and b == 1:
        raise ValueError("Invalid numbers")
    return a - b

### END SOLUTION
```

<!-- grade_id: Question 09, points: 15, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (15)
def exce_sub_(a,b):
    if a == 1 and b == 1:
        raise ValueError("Invalid numbers")
    return a - b
#
#
#
assert exce_sub(2, 1) == exce_sub_(2, 1), "Did you compute the difference correctly?"
assert exce_sub(10,3) == exce_sub_(10,3), "Did you compute the difference correctly?"
assert exce_sub(0, 1) == exce_sub_(0, 1), "Did you compute the difference correctly?"
assert exce_sub(1, 6) == exce_sub_(1, 6), "Did you compute the difference correctly?"
try:
    exce_sub(1, 1)
except Exception as e:
    assert isinstance(e, ValueError), "Make sure to raise a ValueError and not any other type of Exception."
    assert str(e) == "Invalid numbers", "Your error message '{}' is not correct.".format(e)
print("That's correct!") 
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

### Lambda functions

In this final section you will test your knowledge on lambda functions. Remember, you can use the `lambda` keyword to write functions quickly, but only for functions that evaluate to an expression. So this is a good approach for functions you will only use a few times. 

As you may remember, lambda functions look like this. 

```Python
name = lambda argument: what the function will do
```

[Back to top](#Index:) 

### Question 10
*5 points*
  
Create a variable called `f` that stores a lambda function. Your lambda function should take one parameter, `x` and compute the square of it. 

Test your function with the number 6 and store it in a variable called `ans_1`, then test it again with the number 112 and store it in a variable called `ans_2`.

```python
### GRADED

### YOUR SOLUTION HERE
f = None

ans_1 = None
ans_2 = None

### BEGIN SOLUTION
f = lambda x: x*x

ans_1 = f(6)
ans_2 = f(112)
### END SOLUTION
```

<!-- grade_id: Question 10, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (5)
f_ =  lambda x: x*x

ans_1_ = f_(6)
ans_2_ = f_(112)

#
#
#
assert ans_1 == ans_1_, "Did you define your lambda function correctly?"
assert ans_2 == ans_2_, "Did you define your lambda function correctly?"
assert f(10) == f_(10), "Did you define your lambda function correctly?"
assert f(4) == f_(4), "Did you define your lambda function correctly?"
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

[Back to top](#Index:) 

### Question 11
*10 points*
    
You can also define a lambda function that takes more than one argument by separating the arguments by a comma.

Define a lambda function `root` that takes two arguments, `x` and `n`, and computes the `n`th root of x.

Test your function by setting `x = 32` and `n = 5`. 

Assign the result to `ans3`.

Recall, you can compute the `n`th root of a number using the formula 

$$x^{\frac{1}{n}}$$

**HINT:** You can raise a variable to *any* power by using the ** operator.

```python
### GRADED

### YOUR SOLUTION HERE
ans3 = None

### BEGIN SOLUTION
root = lambda x,n: x**(1/n)
ans3 = root(32,5)
### END SOLUTION
```

<!-- grade_id: Question 11, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS (10)
root_ = lambda x,n: x**(1/n)
ans3_ = root_(32,5)
#
#
#
import math
assert math.isclose(ans3, ans3_), "Did you use the correct formula?"
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```

[Back to top](#Index:) 

### Question 12
*10 points*

Define a lambda function `sum_prod` that takes three arguments, `x`, `y` and `n`, that computes sum of `x` and `y` and then multiplies the result by `n` squared. Test your function by setting `x = 14`, `y=45` and `n = 5`. Assign the result to `ans4`.


```python
### GRADED

### YOUR SOLUTION HERE
ans4 = None

### BEGIN SOLUTION
sum_prod = lambda x,y,n: (x+y)*(n**2)
ans4 = sum_prod(14,45,5)
### END SOLUTION
```

```python
### BEGIN HIDDEN TESTS (10)
sum_prod_ = lambda x,y,n: (x+y)*(n**2)
ans4_ = sum_prod_(14,45,5)
#
#
#
assert ans4 == ans4_, "Did you use the correct formula and numbers?"
print("That's correct!")
### END HIDDEN TESTS
```

Output:
```
That's correct!
```